# SIDAS502: Math Methods for Data Science 


# School of Information, University of Michigan


## Week2: Eigenvalue/Eigenvector/Linear Trasformation

### Version 1.2

This assignment is designed to demonstrate basic linear algebra and computational properties of matrices through identifying linearity across vectors, using row reduction to calculate eigenvalues and eigenvectors. Please read the directions carefully, as we want to avoid submissions that are marked incorrect due to formatting mistakes. You will be using sympy, numpy, and scipy for this assignment. 

Ensure that the following code is present at the beginning of your assignment:
- from sympy import *
- import numpy as np
- from scipy import linalg


Please enter your name: "Aseem Sachdeva"

## Part 1: Linearity


<b>1.0</b> In this section, we will be determining if a set of vectors are linearly independent (if they are not, then they are linearly dependent). Decide whether each subset of R^3 is linearly dependent or linearly independent. 

    (1) v1=[-1,3,4],v2=[6,-2,9],v3=[3,8,5] 
    (2) v1=[4,3,3],v2=[6,3,3],v3=[9,3,3] 
    (3) v1=[0,0,-4],v2=[1,-2,9],v3=[3,8,5] 
    (4) v1=[-1,3,6],v2=[6,0,9],v3=[1,8,5]
    


```python
import sympy
from sympy import *
import numpy as np
from scipy import linalg
```

<strong>1.1</strong> \[1 pt\] Using the first set of vectors from (1.0), determine if the vector is linearly independent. Show your work and store your answer "yes" or "no" in a variable called "<strong>linear1</strong>".

<!-- grade_id: cell-411ec698d9ae862d, points: 0, grade: false, solution: true -->
```python
v1=np.array([-1,3,4])
v2=np.array([6,-2,9])
v3=np.array([3,8,5])
linear1="yes"

independent_or_dependent = np.row_stack([v1, v2, v3])

yes_or_no = sympy.Matrix(independent_or_dependent).T.rref()



#raise NotImplementedError()
```

<!-- grade_id: cell-45a28a1ab4ac003c, points: 1, grade: true, solution: false -->
```python
assert linear1.strip().lower() == "yes" or "no", "Problem 1.1, testing linear1, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```

<strong>1.2</strong> \[1 pt\] Using the second set of vectors from (1.0), determine if the vector is linearly independent. Show your work and store your answer "yes" or "no" in a variable called "<strong>linear2</strong>".

<!-- grade_id: cell-9437cf7256c37d72, points: 0, grade: false, solution: true -->
```python
v1=np.array([4,3,3])
v2=np.array([6,3,3])
v3=np.array([9,3,3])
linear2="no"

independent_or_dependent = np.row_stack([v1, v2, v3])

yes_or_no = sympy.Matrix(independent_or_dependent).T.rref()



#raise NotImplementedError()

```

<!-- grade_id: cell-158e70c4cabaca4d, points: 1, grade: true, solution: false -->
```python
assert linear2.strip().lower() == "yes" or "no", "Problem 1.2, testing linear2, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```

<strong>1.3</strong> \[1 pt\] Using the third set of vectors from (1.0), determine if the vector is linearly independent. Show your work and store your answer "yes" or "no" in a variable called "<strong>linear3</strong>".

<!-- grade_id: cell-d68d4ff3c19a01a9, points: 0, grade: false, solution: true -->
```python
v1=np.array([-1,3,6])
v2=np.array([6,0,9])
v3=np.array([1,8,5])
linear3="yes"

independent_or_dependent = np.row_stack([v1, v2, v3])

yes_or_no = sympy.Matrix(independent_or_dependent).T.rref()



#raise NotImplementedError()
```

<!-- grade_id: cell-1b83b01be1cc2c58, points: 1, grade: true, solution: false -->
```python
assert linear3.strip().lower() == "yes" or "no", "Problem 1.3, testing linear3, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```

<strong>1.4</strong> \[1 pt\] Using the fourth set of vectors from (1.0), determine if the vector is linearly independent. Show your work and store your answer "yes" or "no" in a variable called "<strong>linear4</strong>".

<!-- grade_id: cell-61d118295bb6a565, points: 0, grade: false, solution: true -->
```python
v1=np.array([0,0,-4])
v2=np.array([1,-2,9])
v3=np.array([3,8,5])
linear4="yes"

independent_or_dependent = np.row_stack([v1, v2, v3])

yes_or_no = sympy.Matrix(independent_or_dependent).T.rref()





#raise NotImplementedError()
```

<!-- grade_id: cell-0fa75c395d30704c, points: 1, grade: true, solution: false -->
```python
assert linear4.strip().lower() == "yes" or "no", "Problem 1.4, testing linear4, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```

## Part 2: Eigenvalue/ Eigenvector


<strong>2.1</strong> \[1 pt\] What is the eigenvalue and eigenvectors of the matrix[[1,3],[2,5]]? Please store the eigenvalues in the variable called <strong>vals</strong>. The datatype for vals should be an array.



<!-- grade_id: cell-571f59fab4ac61a1, points: 0, grade: false, solution: true -->
```python
v1= np.array([[1,3],[2,5]])

eigen_vals = np.linalg.eigvals(v1)

vals= eigen_vals




#raise NotImplementedError()
```

<!-- grade_id: cell-174c290441138ebc, points: 1, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

<strong>2.2</strong> \[1 pt\] What is the eigenvectors of the matrix[[1,3],[2,5]]? Please store the eigenvectors in a variable called <strong>vecs</strong>. The datatype for vals should be an array.



<!-- grade_id: cell-ae7fca74e80f42d3, points: 0, grade: false, solution: true -->
```python
eigenvalues, eigenvectors = np.linalg.eig(v1)

vecs= eigenvectors






#raise NotImplementedError()
```

<!-- grade_id: cell-c9c31cdf9a7b77ad, points: 1, grade: true, solution: false -->
```python
#hidden tests are within this cell
```

## Part 3: Composition 


<strong>3.1</strong> \[1 pt\] Assume f(x)= 3x+6 and g(y)= 2y+1, what is the reduced equation of h(x)=g * f(x)=g(f(x))? Store your answer in the variable <strong>compo</strong> as a string type.

<!-- grade_id: cell-81395b5bd2a197fb, points: 0, grade: false, solution: true -->
```python
compo="6x + 13"

#raise NotImplementedError()
```

<!-- grade_id: cell-f797c1459bf93874, points: 1, grade: true, solution: false -->
```python
assert type(compo) == str, "Problem 3.1, testing compo, type of value stored in variable does not match the expected type. Expecting String."

```

<strong>3.2</strong> \[1 pt\] Given x = 1, use the equation assigned to compo to solve. Store your answer in the variable <strong>compo1</strong> as an integer. 

<!-- grade_id: cell-beb471e3c613b1ec, points: 0, grade: false, solution: true -->
```python
compo1=19

#raise NotImplementedError()
```

<!-- grade_id: cell-850259d8a0f763ca, points: 1, grade: true, solution: false -->
```python
assert type(compo1) == int, "Problem 3.2, testing compo1, type of value stored in variable does not match the expected type. Expecting Integer."
```

## Part4: Linear Transformation 

<strong>4.1</strong> \[1 pt\] Assume we have an equation: x<sup>2</sup>yz/3+x$\sqrt{y}$+</sup>z+3z =10, is this equation linear in the set of Xs? Store your answer - "yes" or "no" - in a variable called "<strong>ANS41</strong>".

<!-- grade_id: cell-7cd69aa0c7535f5d, points: 0, grade: false, solution: true -->
```python
ANS41="no"

#raise NotImplementedError()
```

<!-- grade_id: cell-4e6661b95f6581b7, points: 1, grade: true, solution: false -->
```python
assert ANS41.strip().lower() == "yes" or "no", "Problem 4.1, testing ANS41, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```

<strong>4.2</strong> \[1 pt\] Assume we have an equation: x<sup>2</sup>yz/3+x$\sqrt{y}$+</sup>z+3z =10, is this equation linear in the set of Ys? Store your answer - "yes" or "no" - in a variable called "<strong>ANS42</strong>".

<!-- grade_id: cell-c2c8d3ec1cbf30a6, points: 0, grade: false, solution: true -->
```python
ANS42="no"

#raise NotImplementedError()
```

<!-- grade_id: cell-489ae00da5e1b5e3, points: 1, grade: true, solution: false -->
```python
assert ANS42.strip().lower() == "yes" or "no", "Problem 4.2, testing ANS42, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```

<strong>4.3</strong> \[1 pt\] Assume we have an equation: x<sup>2</sup>yz/3+x$\sqrt{y}$+</sup>z+3z =10, is this equation linear in the set of Zs? Store your answer - "yes" or "no" - in a variable called "<strong>ANS43</strong>".

<!-- grade_id: cell-b08ce396e53230a1, points: 0, grade: false, solution: true -->
```python
ANS43="yes"

#raise NotImplementedError()
```

<!-- grade_id: cell-a1b171de5ffbdd5c, points: 1, grade: true, solution: false -->
```python
assert ANS43.strip().lower() == "yes" or "no", "Problem 4.3, testing ANS43, value was not 'yes' or 'no'. Please choose 'yes' or 'no' as your answer."
```