# SIDAS502: Math Methods for Data Science 


# School of Information, University of Michigan


## Week4: Probability/Bayes Rule/Maximum likelihood 

### Version 1.3

This assignment is designed to refresh probability skills as well as apply probability in the context of Bayes' Theorem, and maximum likelihood for linear and logistic regression. Please read the directions carefully, as we want to avoid submissions that are marked incorrect due to formatting mistakes. 

Include the following code at the beginning of your assignment:
- from sklearn.model_selection import train_test_split
- from sklearn.linear_model import LogisticRegression
- from sklearn.metrics import classification_report, confusion_matrix
- import pandas as pd
- import numpy as np

Plese enter your name ="Aseem Sachdeva"

```python
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import classification_report, confusion_matrix
import pandas as pd
import numpy as np

```

## 1. Probability

<strong>1.1</strong> \[1 pt\] A college is looking to form a team of 8 to participate in a coding compeition. The team needs at least 2 females and 2 males. If there are 4 males and 7 females who want to join the competition, how many different ways can the team be formed? Please store your answer as a python int (number only, e.g. ANS1=1) into a variable <strong>ANS1</strong>.

<!-- grade_id: cell-ee2bd7ab2781aa26, points: 0, grade: false, solution: true -->
```python
import math
from itertools import combinations
#arr = ['male', 'male', 'male', 'male', 'female', 'female', 'female', 'female', 'female', 'female', 'female']
#r = 8

#x_list = list(combinations(arr,8))



#valid = [i for i in x_list if i.count('male') >= 2 and i.count('female') >= 2 ]

#print(len(x_list))

#print(len(valid))

ANS1= 161
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-2380314f88f56d22, points: 1, grade: true, solution: false -->
```python
assert type(ANS1) == int, "Problem 1.1, testing ANS1, type of value stored in variable does not match the expected type. Expecting Integer"
```

## 2. Conditional Probability

At a school, 40% of the students speak Spanish, 25% of the students speak German, and 15% of the students speak both of the langugues. Now, randomly pick a student. 

<strong>2.1</strong> \[1 pt\] Given that they speak Spanish, what is the probability that they also speak German? Store answer in variable <strong>ANS21</strong> (eg: ANS2a= 0.111).

<!-- grade_id: cell-8566435eee91457e, points: 0, grade: false, solution: true -->
```python
#pB = 0.40
#pAB=0.15

#pAifB = pAB / pB
#print(pAifB)




ANS21 = 0.37499999999999994
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-557a697b7229071a, points: 1, grade: true, solution: false -->
```python
assert type(ANS21) == float, "Problem 2.1, testing ANS21, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>2.2</strong> \[1 pt\] Given that they do not speak German, what is the probability that they speak Spanish? Store answer in variable <strong>ANS22</strong> (eg: ANS2a= 0.11). Please round your answer to the hundredths place (ex. ANS22 = 0.11).

<!-- grade_id: cell-d8a6247556748dcc, points: 0, grade: false, solution: true -->
```python
#P(ppl speaking only Spanish and not German) / P(ppl not speaking German




ANS22=0.33
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-0c3c06d832ee2f42, points: 1, grade: true, solution: false -->
```python
assert type(ANS22) == float, "Problem 2.2, testing ANS22, type of value stored in variable does not match the expected type. Expecting Float."
```

## 3. Independent Probability 

Tom and Lily are trying to solve a problem. The probability of Tom successfully solving the problem is 0.4 and the probability of Lily successfully solving the problem is 0.5. They do not influence each other when they are solving the problem. For each solution, only store the value of probability rounded to the tenth place (ex. ANS = 0.7).

<strong>3.1</strong> \[1 pt\] What is the probability that both of them sucessfully solve the problem? Please store your answer (number only) to a variable <strong>ANS31</strong>.

<!-- grade_id: cell-a89dfc965132afad, points: 0, grade: false, solution: true -->
```python
#p_tom = 0.4
#p_lilly = 0.5

#p_tom_and_lilly = p_tom*p_lilly

#print(p_tom_and_lilly)


ANS31 = 0.20
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-4d4bf277ef89e5a1, points: 1, grade: true, solution: false -->
```python
assert type(ANS31) == float, "Problem 3.1, testing ANS31, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>3.2</strong> \[1 pt\] What is the probability that only one of them successfully solves the problem? Please store your answer (number only) to a variable <strong>ANS32</strong>. 

<!-- grade_id: cell-e55df99dc649e767, points: 0, grade: false, solution: true -->
```python
#p_tom = 0.4
#p_lilly = 0.5
#p_tom_and_lilly = 0.20

#p_only_one = p_tom + p_lilly-(2*p_tom_and_lilly)

#print(p_only_one)

ANS32=0.5
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-04bf980d4b26a9f0, points: 1, grade: true, solution: false -->
```python
assert type(ANS32) == float, "Problem 3.2, testing ANS32, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>3.3</strong> \[1 pt\] Given that **only** one of them successfully solved the problem, what is the probability that the person is Tom?
Please store your answer (number only) to a variable <strong>ANS33</strong>. 

<!-- grade_id: cell-9cd919f33acab2ef, points: 0, grade: false, solution: true -->
```python
#p1⋅(1−p2)+p2(1−p1)=p1+p2−2p1p2.
#p1⋅(1−p2)

#p1 = 0.4
#p2 = 0.5

#test = p1*(1-p2)

#print(test)

###Above code would be for general case. Since these are independent events, answer is just the probability of Tom's event



ANS33 = 0.4
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-9beab75a7dab793d, points: 1, grade: true, solution: false -->
```python
assert type(ANS33) == float, "Problem 3.3, testing ANS33, type of value stored in variable does not match the expected type. Expecting Float."
```

## 4. Bayes Rule

Suppose that you are an art dealer for Van Gogh paintings. Whenever someone brings in a painting that they claim is from Van Gogh, you send in the painting for pigment analysis. Pigment analysis successfully identifies 90% of Van Gogh paintings and falsely identifies 20% of forgeries as being from Van Gogh. In your experience, 99% of the paintings that people bring to you are forgeries.

<strong>4.1</strong> \[1 pt\] Use Bayes theorem to calculate the chance that a painting is from Van Gogh if the pigment analysis finds that it is from Van Gogh. Store your answer in variable <strong>ANS41</strong> as a float (round your answer to the thousandths place, e.g. ANS41=0.111).

<!-- grade_id: cell-f68142d6a4ebf595, points: 0, grade: false, solution: true -->
```python
#P(A|B) = (P(B|A)*PA)/P(B)


ANS41 = 0.043
# YOUR CODE HERE
#raise NotImplementedError()
#print(type(ANS41))
```

<!-- grade_id: cell-c32efd59d0bc94f7, points: 1, grade: true, solution: false -->
```python
assert type(ANS41) == float, "Problem 4.1, testing ANS41, type of value stored in variable does not match the expected type. Expecting Float."
```

## 5. Maximum Likelihood in Linear Regression

Let's say we have serveral data points: (2, 4.5, 8), and we know they are from a normal (Gaussian) distribution. 

<strong>5.1</strong> \[1 pt\] Let's start by assuming some values for mu and sigma. What is the likelihood of guess1: N(mu=4.5,sigma=4)? Store your answer in variable <strong>ANS51</strong> as a float (round your answer to the hundred-thousandths place, e.g.: ANS51 = 0.11111).

```python
import numpy as np
from scipy.stats import norm
```

<!-- grade_id: cell-47ee94371a9ccbf0, points: 0, grade: false, solution: true -->
```python

likelihood = norm.pdf(2, 4.5, 4)*norm.pdf(4.5, 4.5, 4)*norm.pdf(8, 4.5, 4)


likelihood = round(likelihood, 5)


print(likelihood)
ANS51= likelihood

# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
0.00056
```

<!-- grade_id: cell-2cd5d1e205fb8721, points: 1, grade: true, solution: false -->
```python
assert type(ANS51) == np.float64, "Problem 5.1, testing ANS51, type of value stored in variable does not match the expected type. Expecting a numpy float."
```

<strong>5.2</strong> \[1 pt\] Let's change our assumptions regarding the values of mu and sigma. What is the likelihood of guess2: N(mu=6,sigma=8)? Store your answer in variable <strong>ANS52</strong> as a float (round your answer to the hundred-thousandths place, e.g.: ANS52 = 0.11111).

<!-- grade_id: cell-b45ebef5ff62d42b, points: 0, grade: false, solution: true -->
```python

likelihood = norm.pdf(2, 6, 8)*norm.pdf(4.5, 6, 8)*norm.pdf(8, 6, 8)


likelihood = round(likelihood, 5)


print(likelihood)



ANS52 = likelihood 
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
0.0001
```

<!-- grade_id: cell-838756d8041f9a38, points: 1, grade: true, solution: false -->
```python
assert type(ANS52) == np.float64, "Problem 5.2, testing ANS52, type of value stored in variable does not match the expected type. Expecting a numpy float."
```

<strong>5.3</strong> \[1 pt\] Which are the more likely estimates? Store your answer as either "guess1" or "guess2" into a variable called <strong>ANS53</strong>.

<!-- grade_id: cell-2f23f3d7afd2148f, points: 0, grade: false, solution: true -->
```python
ANS53 = "guess 1"
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-4d3946df1270c598, points: 1, grade: true, solution: false -->
```python
assert ((ANS53.strip().replace(" ", "").lower() == "guess1") or (ANS53.strip().replace(" ", "").lower() =="guess2")), "Problem 5.3, testing ANS53, value was not 'guess1' or 'guess2'. Please choose 'guess1' or 'guess2' as your answer."
```

## 6. Likelihood calculation based on a line

We have calculated probability of wins and losses for some logistic regression line. Here are those probabilities: 
- <b>P(win):</b> 0.52, 0.76, 0.89, 0.91, 0.95, 0.95, 0.98, 0.98
- <b>P(loss):</b> 0.08, 0.08, 0.08, 0.15, 0.2, 0.25, 0.6, 0.7


<strong>6.1</strong> \[1 pt\] What is the <strong>likelihood</strong> of the line given the data (you do not need the line nor the data - only the resulting probabilities). Please store your answer into <strong>ANS61</strong>. (Round your answer to the thousandths place) 

<!-- grade_id: cell-de79d768892ac937, points: 0, grade: false, solution: true -->
```python
likelihood = 0.52*0.76*0.89*0.91*0.95*0.95*0.98*0.98*(1-0.08)*(1-0.08)*(1-0.08)*(1-0.15)*(1-0.20)*(1-0.25)*(1-0.60)*(1-0.70)


likelihood = round(likelihood, 3)
print(likelihood)

ANS61=likelihood
# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
0.013
```

<!-- grade_id: cell-4261db1e18760504, points: 1, grade: true, solution: false -->
```python
assert type(ANS61) == float, "Problem 6.1, testing ANS61, type of value stored in variable does not match the expected type. Expecting Float."
assert (ANS61 < 0.00001) == False, "Problem 6.1, testing ANS61, your value is smaller than expected. Did you account for wins/losses?"
```

<strong>6.2</strong> \[1 pt\] What is the <strong>log likelihood</strong> (ln) of the line given the data? Please store your answer into <strong>ANS62</strong> as a float (rounding to the thousandths place, e.g. 1.001).

<!-- grade_id: cell-da056b4724593d40, points: 0, grade: false, solution: true -->
```python
likelihood = np.log(0.52*0.76*0.89*0.91*0.95*0.95*0.98*0.98*(1-0.08)*(1-0.08)*(1-0.08)*(1-0.15)*(1-0.20)*(1-0.25)*(1-0.60)*(1-0.70))


likelihood = round(likelihood, 3)

likelihood.astype(float)

print(likelihood)


ANS62 = -4.326

#raise NotImplementedError()
```

Output:
```
-4.326
```

<!-- grade_id: cell-06fbfe53174dfc2d, points: 1, grade: true, solution: false -->
```python
assert type(ANS62) == float, "Problem 6.2, testing ANS62, type of value stored in variable does not match the expected type. Expecting Float."
assert (ANS62 > 2) == False, "Problem 6.2, testing ANS62, value is larger than expected. Which log did you use (log10 or ln)? Which are you supposed to use?"
assert (ANS62 < -4.5) == False, "Problem 6.2, testing ANS62, value is smaller than expected. Check your math. Did you account for wins/losses?"
```

Sources:
<ul>
<li>https://www.kaggle.com/mashimo/a-very-simple-logistic-regression-model</li>
<li>https://pbpython.com/categorical-encoding.html</li>
<li>https://www.kaggle.com/chethan187/classification-model-using-logistic-regression</li>
</ul>